<a href="https://colab.research.google.com/github/shahreeza-kassim/WID3003-NeuralComputing/blob/main/03_ForwardAD_question_upload.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Forward Mode Automatic Differentiation (AD)

Forward mode AD can simply be implemented by defining a class to represent [dual numbers](https://en.wikipedia.org/wiki/Dual_number) which hold the value and its derivative. The following code defines a dual number and implements multiplication. 

Tasks:

1. Addition (`__add__`) is incomplete - Please complete it 
2. Implement division (`__truediv__`), subtraction (`__sub__`) and power (`__pow__`) as well

<!-- grade_id: cell-a645b47036f56c2e, points: 0, grade: false, solution: true -->
```python
import math

class DualNumber:
    def __init__(self, value, dvalue): 
        self.value = value
        self.dvalue = dvalue #derivative

    def __str__(self):
        return str(self.value) + " + " + str(self.dvalue) + "ε"

    def __mul__(self, other):  #multiplication  udv+vdu
        return DualNumber(self.value * other.value,
            self.dvalue * other.value + other.dvalue * self.value)
    
    def __add__(self, other): 
        #Your code here
        raise NotImplementedError()
    

```

Output:
```
6 + 10ε
```

<!-- grade_id: cell-7923d5c39d574cd7, points: 3, grade: true, solution: false -->
```python
# Tests

result = DualNumber(1,0) + DualNumber(1,0) / DualNumber(1,0) - DualNumber(1,0)**DualNumber(1,0)


```

## Implementing math functions

We also need to implement some core math functions. Here's the sine function for a dual number:

```python
def sin(x):
    return DualNumber(math.sin(x.value), math.cos(x.value)*x.dvalue)
```

Task: Implement the _cosine_ (`cos`), _tangent_ (`tan`), and _exponential_ (`exp`) functions in the code block below.

<!-- grade_id: cell-61b9462cb781a8e1, points: 0, grade: false, solution: true -->
```python
# Implement additional math functions on dual numbers

def cos(x):
    raise NotImplementedError()    

def tan(x):
    raise NotImplementedError()    

def exp(x):
    raise NotImplementedError()
```

<!-- grade_id: cell-9055ec23e6e9ffb8, points: 3, grade: true, solution: false -->
```python
# Tests
assert cos(DualNumber(0,0)).value == 1
assert tan(DualNumber(0,0)).value == 0
assert exp(DualNumber(0,0)).value == 1

'''
The assert keyword lets you test if a condition in your code returns True, 
if not, the program will raise an AssertionError
'''

```

## Putting it all together

1. Write a code to compute the value of the function $z=x\cdot y+sin(x)$ given $x=0.5$ and $y=4.2$, together with the derivative $\partial z/\partial x$ at that point. 



<!-- grade_id: cell-12f7c8fa42784946, points: 0, grade: false, solution: true -->
```python



```

2. Differentiate the above function with respect to $x$ and write the symbolic derivatives. Verify the result computed above is correct by plugging-in the values into your symbolic gradient expression.

3. Compute the derivative $\partial z/\partial y$ of the above expression (at the same point $x=0.5, y=4.2$ as above). Verify by hand that the result is correct.

4. Experiment with other math functions and methods you created